Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Matthew Sah"
COLLABORATORS = "" 
```

---

# Problem 1 (3 points)
Write a function to calcuate the area of a triangle given the length of its three sides--a, b, and c--using these formulas: $s=(a+b+c)/2$ $A=\sqrt{s(s-a)(s-b)(s-c)}$.

```python
import math
```

<!-- grade_id: calculate_area_triange, points: 0, grade: false, solution: true -->
```python
def calculate_area_triange(a, b, c): 
    S = ( a + b + c ) / 2
    A = round(float(math.sqrt(S * (S-a) * (S-b) * (S-c))), 2) 
    return A 
```

<!-- grade_id: calculate_area_triange_test, points: 3, grade: true, solution: false -->
```python
assert calculate_area_triange(3,4,5) == 6.0
assert calculate_area_triange(10,15,20) == 72.62
assert calculate_area_triange(3.5,4.5,7.5) == 5.17
```

# Problem 2 (5 points)
A certain CS professor gives 5-point quizzes that are graded on the scale 5-A, 4-B, 3-C, 2-D, 1-F, 0-5. Write a function that accepts a quiz score as an input and returns the corresponding grade. 

<!-- grade_id: calculate_grade, points: 0, grade: false, solution: true -->
```python
def calculate_grade_letter(score):
    
    if score == 0:
        return "F"
    elif score == 1:
        return "F"
    elif score == 2:
        return "D"
    elif score == 3:
        return "C"
    elif score == 4:
        return "B"
    elif score == 5:
        return "A"
    return "F" 
```

<!-- grade_id: calculate_grade_test, points: 5, grade: true, solution: false -->
```python
assert calculate_grade_letter(0) == "F"
assert calculate_grade_letter(1) == "F"
assert calculate_grade_letter(2) == "D"
assert calculate_grade_letter(3) == "C"
assert calculate_grade_letter(4) == "B"
assert calculate_grade_letter(5) == "A"
```

# Problem 3 (5 points)
A certain CS professor gives 100-point exams that are graded on the scale 90-100:A, 80-89:B, 70-79:C, 60-69:D, <60:F.
Write a function that accepts an exam score as input and returns the corresponding grade. 

<!-- grade_id: calculate_grade_letter2, points: 0, grade: false, solution: true -->
```python
def calculate_grade_letter2(score):
    
    if score <= 100  and score >= 90:
        return "A"
    elif score <= 89 and score >= 80:
        return "B"
    elif score <= 79 and score >= 70:
        return "C"
    elif score <= 69 and score >= 60:
        return "D"
    elif score <= 59 and score >= 0 :
        return "F"
    return null
     
```

<!-- grade_id: calculate_grade_letter2_test, points: 5, grade: true, solution: false -->
```python
assert calculate_grade_letter2(95) == "A"
assert calculate_grade_letter2(88) == "B"
assert calculate_grade_letter2(73) == "C"
assert calculate_grade_letter2(64) == "D"
assert calculate_grade_letter2(55) == "F"

```

# Problem 4 (3 points)
Write a function that counts the number of words in a sentence entered by the user.

<!-- grade_id: count_words_in_sentence, points: 0, grade: false, solution: true -->
```python
def count_words_in_sentence(sentence): 
    return len(sentence.split()) 
```

<!-- grade_id: count_words_in_sentence_test, points: 3, grade: true, solution: false -->
```python
assert count_words_in_sentence("This sentence has five words!") == 5
assert count_words_in_sentence("Function takes in a sentence and returns the numbers of words in the sentence") == 14
assert count_words_in_sentence("Function takes in as input three sides of a triangle, a, b, and c, calculates its area and returns the area. Round the output to two decimal places.") == 28

```

# Problem 5 (3 points)
Wirte a function that calculates the average word length in a sentence entered by the user.

<!-- grade_id: calculate_average_word_length, points: 0, grade: false, solution: true -->
```python
def calculate_average_word_length(sentence):
    words = sentence.split()
    average = round(sum(len(word) for word in words) / len(words),2)
    print(average)
    return average 
```

<!-- grade_id: calculate_average_word_length_test, points: 3, grade: true, solution: false -->
```python
assert calculate_average_word_length("This sentence has five words!") == 5.0
assert calculate_average_word_length("Function takes in a sentence and returns the numbers of words in the sentence") == 4.57
assert calculate_average_word_length("Function takes in as input three sides of a triangle, a, b, and c, calculates its area and returns the area. Round the output to two decimal places.") == 4.32

```

Output:
```
5.0
4.57
4.32
```

# Problem 6 (3 points)
Write a function that takes in a list and modifies the list by squaring each entry.

<!-- grade_id: square_each, points: 0, grade: false, solution: true -->
```python
def square_each(nums_list):
    returnList = []
    for i in range(0, len(nums_list)): 
        nums_list[i] = nums_list[i] * nums_list[i] 
    return returnList 
```

<!-- grade_id: square_each_test, points: 3, grade: true, solution: false -->
```python
nums_list = [1,2,3]
square_each(nums_list)
assert nums_list == [1, 4, 9]

nums_list = [6,5,4]
square_each(nums_list)
assert nums_list == [36, 25, 16]


nums_list = [2.5,3.5,4.5]
square_each(nums_list)
assert nums_list == [6.25, 12.25, 20.25]


```