Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Matthew Sah"
COLLABORATORS = ""
```

---

# Problem 1 (7 points)

The body mass index (BMI) is calcualted as a person's weight (in pounds)
times 720, divided by the square of the person's height (in inches). A BMI
in the range 19-25, inclusive, is considered healthy. Write a function
that calculates a person's BMI and RETURNS one of three string: 'above', 'within',
or 'below'. Note the function takes as input the person's weight and height.
This means you do not have to prompt the user for the input. Rather, the input is
provided when the function is called.

The function MUST raise exception type ValueError if the weight or height are less
than or equal to zero.


<!-- grade_id: problem_1, points: 0, grade: false, solution: true -->
```python
def calculate_bmi(weight, height): 
    if weight <=0 or height <= 0:
        raise ValueError("Values are less than or equal to zero")
    else: 
        BMI = weight * 720 / (height * height )
        
        if BMI > 25:
            return("above")
        elif BMI <= 25 and BMI >=19:
            return("within")
        elif BMI < 19:
            return("below")
     
    
    
```

<!-- grade_id: problem_1_tests, points: 7, grade: true, solution: false -->
```python
assert calculate_bmi(160, 5.6*12) == 'above'
assert calculate_bmi(110, 5.6*12) == 'below'
assert calculate_bmi(135, 5.6*12) == 'within'


try:
    calculate_bmi(160, 0)
except ValueError as error:
    assert error.__class__.__name__ == 'ValueError'
try:
    calculate_bmi(0, 5.5)
except ValueError as error:
    assert error.__class__.__name__ == 'ValueError'

try:
    calculate_bmi(160, -10)
except ValueError as error:
    assert error.__class__.__name__ == 'ValueError'
try:
    calculate_bmi(-10, 5.5)
except ValueError as error:
    assert error.__class__.__name__ == 'ValueError'
```

# Problem 2 (5 points)
Given a person's height, write a function that returns the healthy BMI range.
The return value should be a tuple. Round up the lowest value and round down the
highest value.

The function MUST raise a ValueError type exception if the weight or height are less
than or equal to zero.

<!-- grade_id: problem_2, points: 0, grade: false, solution: true -->
```python
def calculate_within_bmi_range(height): 
    import math  
    BMI_High = round((25 * height * height) / 720)
    BMI_Low = round((19 * height * height) / 720) 
    if BMI_Low <=0 or BMI_High <= 0 or height <=0:
        raise ValueError("Values are less than or equal to zero")
    return int(BMI_Low), int(BMI_High)
```

<!-- grade_id: problem_2_tests, points: 5, grade: true, solution: false -->
```python
assert calculate_within_bmi_range(5.0*12) == (95, 125)
assert calculate_within_bmi_range(5.5*12) == (115, 151)
assert calculate_within_bmi_range(6.0*12) == (137, 180)

try:
    calculate_within_bmi_range(0)
except ValueError as error:
    assert error.__class__.__name__ == 'ValueError'
try:
    calculate_within_bmi_range(-10)
except ValueError as error:
    assert error.__class__.__name__ == 'ValueError'

```

# Problem 3 (4 points)

A year is a leap year if it is divisible by 4, unless it is a century year
that is not divisible by 400. (1800 and 1900 are not leap years while 1600
and 2000 are.) Write a function that determines if a given a year is a leap
year. The function should return True or False.


<!-- grade_id: problem_3, points: 0, grade: false, solution: true -->
```python
def is_leap_year(year):
    if year %4 == 0 and year % 100 ==0  and year %400 ==0:
        return True 
    elif year %4 == 0 and year % 100 != 0:
        return True
    else :
        return False
```

<!-- grade_id: problem_3_tests, points: 4, grade: true, solution: false -->
```python
assert is_leap_year(1775) == False
assert is_leap_year(1776) == True
assert is_leap_year(1800) == False
assert is_leap_year(2000) == True

```

# Problem 4 (4 points)

Given a date in the form month/day/year, write a function that returns
True or False if the date is valid or invalid, respectively. For example
5/24/1962 is valid, but 9/31/2000 is invalid (September has only 30 days).


<!-- grade_id: problem_4, points: 0, grade: false, solution: true -->
```python
def is_date_valid(date):
    parsed = date.split("/")
    import datetime 
    try :
        datetime.datetime(year = int(parsed[2]) ,month =  int(parsed[0]), day = int(parsed[1])) 
        return True
    except:
        return False
```

<!-- grade_id: problem4_tests, points: 4, grade: true, solution: false -->
```python
assert is_date_valid('2/29/1800') == False
assert is_date_valid('5/24/1962') == True
assert is_date_valid('2/29/2000') == True
assert is_date_valid('9/31/2000') == False

```

# Problem 5 (3 points)

Write a function that uses a while toop to determine how long it takes
for an investment to double at a given interest rate. The function
should return the number of years it takes an investment to double.

The interest rate is simple annual and is given as a percentage.
See https://www.investopedia.com/terms/f/futurevalue.asp for formula


<!-- grade_id: problem_5, points: 0, grade: false, solution: true -->
```python
def years_to_2x_investment_simple(initial_investment, interest_rate):
    total_years = 0 
    future_value = 0
    threshold = initial_investment * 2
    while future_value < threshold:
        total_years += 1
        rate_per_year = ( 1 + (interest_rate * total_years * 0.01))
        future_value = initial_investment * rate_per_year 
    print(int(total_years))
    return int(total_years)
```

<!-- grade_id: problem_5_tests, points: 3, grade: true, solution: false -->
```python
assert years_to_2x_investment_simple(100000, 5) == 20
assert years_to_2x_investment_simple(100000, 10) == 10
assert years_to_2x_investment_simple(100000, 15) == 7

```

Output:
```
20
10
7
```

# Problem 6 (3 points)

Similar to problem 5, but the interest rate is compounded annual.

Write a function that uses a while toop to determine how long it takes
for an investment to double at a given interest rate. The function
should return the number of years it takes an investment to double.

The interest rate is COMPOUND annual and is given as a percentage.
See https://www.investopedia.com/terms/f/futurevalue.asp for formula


<!-- grade_id: problem_6, points: 0, grade: false, solution: true -->
```python
def years_to_2x_investment_compound(initial_investment, interest_rate):
    total_years = 0 
    future_value = 0
    threshold = initial_investment * 2
    while future_value < threshold:
        total_years+=1 
        rate_per_year = ( 1 + (interest_rate * 0.01)) ** total_years 
        future_value = initial_investment * rate_per_year   
    return int(total_years)
```

<!-- grade_id: problem_6_tests, points: 3, grade: true, solution: false -->
```python
assert years_to_2x_investment_compound(100000, 5) == 15
assert years_to_2x_investment_compound(100000, 10) == 8
assert years_to_2x_investment_compound(100000, 15) == 5
```