Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Matthew Sah"
COLLABORATORS = ""
```

---

## Part A (3 points)

Write a function that converts tempratues from Fahrenheit to Celsius

<!-- grade_id: fahrenheit_to_celsius, points: 0, grade: false, solution: true -->
```python
def fahrenheit_to_celsius(temprature_in_fahrenheit): 
    return (temprature_in_fahrenheit-32 ) * 5 / 9
    raise NotImplementedError()
```

<!-- grade_id: correct_fahrenheit_to_celsius, points: 3, grade: true, solution: false -->
```python
assert fahrenheit_to_celsius(212) == 100
assert fahrenheit_to_celsius(32) == 0
assert round(fahrenheit_to_celsius(73.4)) == 23.0
```

## Part B (1 points)

Write a function that prompts the user using the `input` function to enter temprature in Fahrenheit. This function should use your `fahrenheit_to_celsius` function. It should print the converted value.

<!-- grade_id: prompt_user_for_temperature_in_fahrenheit, points: 1, grade: true, solution: true -->
```python
def prompt_user_for_temperature_in_fahrenheit():
    """
    Function uses the `input` function to prompt user to enter temperature in Fahrenheit. The function will print
    the converted value. 
    """
    # YOUR CODE HERE
    temperaturer = int(input("Enter Fahrenheit : "))
    return fahrenheit_to_celsius(temperaturer)
    raise NotImplementedError()
```

## Part C (3 points)

Write a function that converts distances measured in kilometers to miles. One kilometer is approximately 0.62 miles.

<!-- grade_id: km_to_miles, points: 0, grade: false, solution: true -->
```python
def km_to_miles(distance_in_km):
    """
    Write a function that converts distances measured in kilometers to miles. 
    One kilometer is approximately 0.62 miles.
    """
    # YOUR CODE HERE 
    return distance_in_km * 0.62
    raise NotImplementedError()
```

<!-- grade_id: correct_km_to_miles, points: 3, grade: true, solution: false -->
```python
assert km_to_miles(160) == 99.2
assert km_to_miles(40) == 24.8
assert km_to_miles(80) == 49.6
```

## Part D (1 points)

Write a function that prompts the user using the `input` function to enter distance in kilometers. This function should use your `km_to_miles` function. It should print the converted value.

<!-- grade_id: prompt_user_for_distance_in_kilometers, points: 1, grade: true, solution: true -->
```python
def prompt_user_for_distance_in_kilometers():
    """
    Function uses the `input` function to prompt user to enter distance in kilometers. The function will print
    the converted value. 
    """
    # YOUR CODE HERE 
    kilometer = int(input("Enter Kilometer : "))
    return km_to_miles(kilometer)
    raise NotImplementedError()
```

## Part E (3 points)

Write a function that calculates the volume of a sphere given its radius. The formula for volume is: $V = 4/3 \pi r^3$. Use $\pi=3.14$. Round the volume to two decimal places. Use the `round` function, e.g., `round(3.1415926535,2)`.

<!-- grade_id: calculate_sphere_volume, points: 0, grade: false, solution: true -->
```python
def calculate_sphere_volume(radius):
    """
    Function calcuates the volume of sphere given its radius
    """
    # YOUR CODE HERE 
    return round(4 / 3 * 3.14 * radius * radius * radius, 2)
    raise NotImplementedError()
```

<!-- grade_id: correct_calculate_sphere_volume, points: 3, grade: true, solution: false -->
```python
assert calculate_sphere_volume(10) == 4186.67
assert calculate_sphere_volume(35) == 179503.33
assert calculate_sphere_volume(78) == 1986791.04
```

## Part F (3 points)

Write a function that calculates the surface area of a sphere given its radius. The formula for surface area is: $A = 4\pi r^2$. Use $\pi=3.14$. Round the surface area to two decimal places. Use the `round` function, e.g., `round(3.1415926535,2)`.

<!-- grade_id: calculate_sphere_surface_area, points: 0, grade: false, solution: true -->
```python
def calculate_sphere_surface_area(radius):
    """
    Function calcuates the surface area of sphere given its radius
    """
    # YOUR CODE HERE 
    return round( 4 * 3.14 * radius * radius , 2)
    raise NotImplementedError()
```

<!-- grade_id: correct_calculate_sphere_surface_area, points: 3, grade: true, solution: false -->
```python
assert calculate_sphere_surface_area(10) == 1256.0
assert calculate_sphere_surface_area(35) == 15386.0
assert calculate_sphere_surface_area(78) == 76415.04
```